# Teste

Antes da submissão final do seu problema, certifique que tudo funciona como esperado. Primeiro, **reinicie o kernel** (Kernel$\rightarrow$Restart) e depois **execute todas as células** (Cell$\rightarrow$Run All).

Preencha todos lugares que digam `SUA SOLUÇÃO` or `YOUR ANSWER HERE`, e o seu nome e dos colaboradores (outros alunos no grupo) abaixo.

```python
# Instale o maestro, se já não estiver instalado
!pip install toolz --user
!pip install sidekick  --user
!pip install hyperpython  --user
!pip install maestro-tools==0.1.0 --user
```

```python
# Digite o nome de cada aluno no formato "Nome (matrícula)"
#
# Se você estiver resolvendo a prova sozinho(a), 
# deixe COLABORATORS como uma string vazia.

NAME = "Nome 1 (1234567)"
COLLABORATORS = ""
```

```python
# Verifica se o nome/matrícula é válido
import re

regex = re.compile(r'[^\)]+(\(.*\))|')
assert all(regex.fullmatch(x) for x in [NAME, COLLABORATORS]), 'Nome incorreto!'
```

---

```python
import re
from maestro.course.compilers import (
    CROSSWORD_EXAMPLES as CW, 
    regex_crossword_table, 
    check_regex_crossword, 
    check_regex, 
)
```

---
## Regex: abab's (1 pt)

Salve uma expressão regular na variável `regex` que reconheça strings formadas por sequências alternadas de a's e b's iniciadas em a.

```python
# Lista de casos de teste (não edite!)
examples = """
POSITIVE:
a
ab
aba
abababa

NEGATIVE:
aa
bb
abba
abc
ba
baba
b
"""
```

<!-- grade_id: ababs-answer, points: 0, grade: false, solution: true -->
```python
# Sua resposta aqui!
regex = r"<sua regex>"
...
```

<!-- grade_id: ababs-test, points: 0.75, grade: true, solution: false -->
```python
"Verifica a expressão regular (75% da nota)"

check_regex(regex, examples)
```

<!-- grade_id: ababs-size, points: 0.25, grade: true, solution: false -->
```python
"Identifica se resolveu com o número mínimo de caracteres (25% da nota)"

check_regex(regex, examples, max_size=8)
```

---
## Regex: abab's II (1 pt)

Salve uma expressão regular na variável `regex` que reconheça strings formadas por sequências alternadas de a's e b's.

```python
# Lista de casos de teste (não edite!)
examples = """
POSITIVE:
a
ab
aba
abababa
b
ba
bab
baba

NEGATIVE:
aa
bb
abba
abc
"""
```

<!-- grade_id: ababs2-answer, points: 0, grade: false, solution: true -->
```python
# Sua resposta aqui!
regex = r"<sua regex>"
...
```

<!-- grade_id: ababs2-test, points: 0.75, grade: true, solution: false -->
```python
"Verifica a expressão regular (75% da nota)"

check_regex(regex, examples)
```

<!-- grade_id: ababs2-size, points: 0.25, grade: true, solution: false -->
```python
"Identifica se resolveu com o número mínimo de caracteres (25% da nota)"

check_regex(regex, examples, max_size=15)
```

---
## Regex: hexadecimais (0.5 pts)

Crie uma expressão regular que identifique números hexadecimais no formato 0xbeef01 ou 
0xF00BA4.

```python
# Lista de casos de teste (não edite!)
examples = """
POSITIVE:

0x123
0x49
0xa
0xA
0xBeef
0xF00BA4

NEGATIVE:

0xFOO
0xbar
123
42
Beef
0x
0
"""
```

<!-- grade_id: hex-answer, points: 0, grade: false, solution: true -->
```python
# Sua resposta aqui!
regex = r"<sua regex>"
...
```

<!-- grade_id: hex, points: 0.25, grade: true, solution: false -->
```python
"Verifica a expressão regular (50% da nota)"

check_regex(regex, examples)
```

<!-- grade_id: hex-size, points: 0.25, grade: true, solution: false -->
```python
"Identifica se resolveu com o número mínimo de caracteres (50% da nota)"

check_regex(regex, examples, max_size=13)
```

---
## Regex: hexadecimais com underscore (1.5 pts)

Crie uma expressão regular que identifique números hexadecimais no formato 0xbeef_01 ou 
0xF00_BA4 onde os underscores são aceitos entre qualquer dígito hexadecimal, mas não podem 
aparecer em sequência ou no início ou fim do número.

```python
# Lista de casos de teste (não edite!)
examples = """
POSITIVE:

0x1_2_3
0x49
0xAa
0x600d
0xF00_BA4

NEGATIVE:

0xbar
42
Beef
0x
0x_123
0x123_
0x600d__123
"""
```

<!-- grade_id: h_e_x-answer, points: 0, grade: false, solution: true -->
```python
# Sua resposta aqui!
regex = r"<sua regex>"
...
```

<!-- grade_id: h_e_x, points: 1.125, grade: true, solution: false -->
```python
"Verifica a expressão regular (75% da nota)"

check_regex(regex, examples)
```

<!-- grade_id: h_e_x-size, points: 0.375, grade: true, solution: false -->
```python
"Identifica se resolveu com o número mínimo de caracteres (25% da nota)"

check_regex(regex, examples, max_size=28)
```

---
## Palavras cruzadas Regex (3pts)

Resolva a palavra cruzada de expressões regulares abaixo. A palavra cruzada de expressões regulares funciona de forma parecida com uma palavra cruzada tradicional, mas utiliza expressões regulares para definir o conteúdo de cada linha e cada coluna. 

Vocë deve completar cada célula na palavra cruzada com uma única letra. A palavra cruzada estará correta se todas linhas e todas colunas satisfizerem as expressões regulares correspondentes.

É um jogo que valoriza a lógica e o raciocínio ao invés da memorização. Complete as duas palavras cruzadas abaixo e salve a resposta em uma varável chamada "answer", que pode ser uma string onde cada linha corresponde
a uma linha da resposta.

---
## Procurando padrões (3pts)

O texto abaixo possui vários números que representam grandezas físicas. Crie uma função que identifica e soma os valores de todos os números de uma determinada grandeza física e retorna a soma dos mesmos.

```
Massa: 42kg
Distâncias: 1.51km, -3.21km, +10km
Duração: 5.0s
```

Os números podem ter parte decimal e sinais opcionais. A grandeza física deve seguir o número sem qualquer espaço e consiste apenas em letras maiúsculas ou minúsculas.

Utilize o esqueleto fornecido abaixo e não altere a assinatura da função.

**Dica**: Digite `re.finditer?` em uma célula para obter a ajuda da função finditer de
expressões regulares.

<!-- grade_id: cell-d1209f9b9e4fa066, points: 0, grade: false, solution: true -->
```python
number_re = re.compile(r'([+-]?\d+(?:\.\d+)?)([a-zA-Z]+)')

def sum_values(text, kind):
    ...
```

```python
# Área de experimentação. Faça o que você quiser aqui

text = """
Massa: 42kg
Distâncias: 1.51km, -3.21km, +10km
Duração: 5.0s
"""
sum_values(text, "km")
```

<!-- grade_id: cell-6e1db37cde6b2c07, points: 0.5, grade: true, solution: false -->
```python
"Verifica inteiros"

assert sum_values("1m, +2m, -4m", "m") == -1
```

<!-- grade_id: cell-e173ad18149803fa, points: 0.75, grade: true, solution: false -->
```python
"Verifica se filtra grandezas"

assert sum_values("1.0m, 2.0m, 3.0s, 4.0m, 5", "m") == 7.0
```

<!-- grade_id: cell-239f7170f4b285a6, points: 0.75, grade: true, solution: false -->
```python
"Utiliza grandezas válidas"

assert sum_values("1.125m1m", "m") == 2.125
```

<!-- grade_id: phys-ex, points: 1, grade: true, solution: false -->
```python
"Verifica texto de exemplo"

text = """
Massa: 42kg
Distâncias: 1.51km, -3.21km, +100km
Duração: 5.0s
"""

assert sum_values(text, 'km') == 98.3
```

---
## Palavras cruzadas Regex (3pts)

Resolva a palavra cruzada de expressões regulares abaixo. A palavra cruzada de expressões regulares funciona de forma parecida com uma palavra cruzada tradicional, mas utiliza expressões regulares para definir o conteúdo de cada linha e cada coluna. 

Vocë deve completar cada célula na palavra cruzada com uma única letra. A palavra cruzada estará correta se todas linhas e todas colunas satisfizerem as expressões regulares correspondentes.

É um jogo que valoriza a lógica e o raciocínio ao invés da memorização. Complete as duas palavras cruzadas abaixo e salve a resposta em uma varável chamada "answer", que pode ser uma string onde cada linha corresponde
a uma linha da resposta.

### Palavra cruzada (yr, 1pt)

```python
regex_crossword_table(*CW["yr"])
```

<!-- grade_id: cw1-answer, points: 0, grade: false, solution: true -->
```python
# Coloque sua resposta aqui!
answer = """
linha-1
linha-2
linha-3
"""
...
```

```python
# Área de experimentação (utilize como você quiser)
...
```

<!-- grade_id: cw1, points: 1.5, grade: true, solution: false -->
```python
"Verifica se a resposta está correta"

check_regex_crossword(answer, *CW["yr"])
```

### Palavra cruzada (john, 1pt)

```python
regex_crossword_table(*CW["john"])
```

<!-- grade_id: cw2-answer, points: 0, grade: false, solution: true -->
```python
# Coloque sua resposta aqui!
answer = """
linha-1
linha-2
linha-3
"""
...
```

```python
# Área de experimentação (utilize como você quiser)
...
```

<!-- grade_id: cw2, points: 1.5, grade: true, solution: false -->
```python
"Verifica se a resposta está correta"

check_regex_crossword(answer, *CW["john"])
```